# Chapter 1 Dynamic Programming and Bellman Equations


## Example: Secretary Problem

In this example, we will solve the secretary problem using stochastic dynamic programming (DP).

### Secretary Problem

The secretary problem [1] can be described as follows. Imagine an administrator who wants to hire the best secretary out of $n$ available candidates. They are interviewed in succession in random order, that is, for each interview the administrator picks a candidate uniformly at random from the remaining candidates. After interviewing the $h^{\mathrm{th}}$ candidate, the administrator knows the ranking of the $h^{\mathrm{th}}$ candidate among the candidates interviewed so far, but is unaware of the quality of unseen candidates. The administrator needs to determine whether to hire or reject the $h^{\mathrm{th}}$ candidate immediately after the interview. Note that the administrator cannot go back to hire candidate $1,\ldots,h-1$ after interviewing the $h^{\mathrm{th}}$ candidate. At most one candidate can be hired. The goal is to maximize the probability of hiring the best candidate.

### Problem Formulation

We can formulate the secretary problem into a finite-horizon stochastic DP problem as follows:

- Stage $h$: the time when the administrator makes decision after interviewing the $h^{\mathrm{th}}$ candidate. $h\in\{1,...,n\}$.
    
- State $s_h$: the ranking of the $h^{\mathrm{th}}$ candidate among the candidates interviewed so far. $s_h\in\{1,...,h\}$.

    Define a terminal state $g$ such that if the administrator has already hired one candidate in one of the previous interviews, then $s_{h}=g, s_{h+1}=g,\ldots,s_{n}=g$. That is, once the administrator decides to hire a candidate, the process will go to the terminal state and remain in the terminal state.

- Action $a_h$:
    
  If $s_h\neq g$, then

  - $a_h=1$ means hiring the $h^{\mathrm{th}}$ candidate;
  - $a_h=0$ means rejecting the $h^{\mathrm{th}}$ candidate.

  If $s_h=g$, then the only available action is $a_h=0$, meaning rejecting the $h^{\mathrm{th}}$ candidate, because $s_h=g$ means that the administrator has already hired one candidate previously.
    
- Reward $r_h(s_h,a_h)$: 
    
    $r_h(s_h,a_h)=1$ if $s_h\neq g$, $a_h=1$ and the $h^{\mathrm{th}}$ candidate is the best candidate among all the $n$ candidates. 
    
    Otherwise, $r_h(s_h,a_h)=0$.
    
    Note that if $s_h\neq 1$, $r_h(s_h,1)=0$ because $s_h\neq 1$ means that the $h^{\mathrm{th}}$ candidate is not the best among the $h$ candidates that have been interviewed and of course not the best among all candidates. If $s_h=1$, we can calculate the probability that the $h^{\mathrm{th}}$ candidate is the best among all candidates as follows:

    \begin{align*}
        \Pr\left(\mbox{the $h^{\mathrm{th}}$ candidate is the best among all candidates}|s_h=1\right)=\frac{{n-1 \choose h-1} (h-1)!(n-h)!}{{n \choose h} (h-1)!(n-h)!}=\frac{h}{n}.
    \end{align*}

    Therefore, $r_h(1,1)=1$ with probability $h/n$.
    
- Transition probabilities:
    
    For any $m\in\{1,\ldots,h+1\}$ and $k\in\{1,\ldots,h\}$, we have

    \begin{align*}
        \Pr(s_{h+1}=m|s_h=k,a_h=0)=\frac{{n \choose h+1} (h-1)!(n-h-1)!}{{n \choose h+1} {h+1 \choose h}(h-1)!(n-h-1)!}=\frac{1}{h+1},
    \end{align*}

    which means that $s_{h+1}$ will be uniformly distributed in $\{1,\ldots,h+1\}$ if $s_h\neq g$ and $a_h=0$. And for any $k\in\{1,\ldots,h\}$,

    \begin{align*}
        &\Pr(s_{h+1}=g|s_h=k,a_h=0)=0,\\
        &\Pr(s_{h+1}=g|s_h=k,a_h=1)=1,\\
        &\Pr(s_{h+1}=g|s_h=g,a_h=0)=1.
    \end{align*}
    
- Goal: maximize the probability of hiring the best candidate. Note that this goal can be written as 
    \begin{align*}
        \max_{\mu_1,\ldots,\mu_n} \mathbb{E}\left[\sum_{h=1}^{n} r_h(s_h,\mu_h(s_h))\right].
    \end{align*}

### Value Function

Let $V_h(s_h)$ denote the optimal value function for state $s_h$ at stage $h$, defined by
$$
V_h(s_h) = \max_{\mu_h,\ldots,\mu_n} \mathbb{E} \left[\sum_{k=h}^n r_k(s_k, \mu_k(s_k))\right].
$$
The optimal value function $V_h(s_h)$ can be interpreted as the probability that the best candidate will be hired under the optimal policy given $s_h$ at stage $h$. 

For example, let $n = 10$, $h = 5$, $s_5 = 1$. In this scenario, there are 10 candidates total. After interviewing the $5^{\mathrm{th}}$ candidate, we know that this candidate ranks first among the 5 candidates that have been interviewed. Then $V_5(1)$ is the probability that we can successfully hire the best candidate if we follow the optimal policies at and after stage $5$.

### Bellman Equation

Based on the above definition and analysis, the Bellman equation for $h=1,...,n-1$ can be written as:

\begin{equation}
    \begin{aligned}
        V_h(1) =& \max \left\{\mathbb{E}[r_h(1,1)], \mathbb{E}[V_{h+1}(s_{h+1})]\right\} = \max\left\{\frac{h}{n},\frac{1}{h+1}\sum_{s=1}^{h+1}V_{h+1}(s)\right\}\\
        V_h(s_{h}) =& \max \left\{0, \mathbb{E}[V_{h+1}(s_{h+1})]\right\} = \frac{1}{h+1}\sum_{s=1}^{h+1}V_{h+1}(s), s_h=2,...,h.
    \end{aligned}
\end{equation}

For $V_h(1)$, the first term inside the maximum function, $\mathbb{E}[r_h(1,1)]$, is the expected reward given $a_h=1$, i.e., hiring the $h^{\mathrm{th}}$ candidate. The total future reward is 0 since the process terminates. The second term $\mathbb{E}[V_{h+1}(s_{h+1})]$ is the expected value given $a_h=0$, i.e., rejecting the $h^{\mathrm{th}}$ candidate, since the instantaneous reward is 0.

For $V_h(s_{h})$ where $s_h=2,...,h$, the expected value for hiring the $h^{\mathrm{th}}$ candidate is 0 since the ranking of the $h^{\mathrm{th}}$ candidate among all the $n$ candidates must be larger than or equal to 2. Similar to the case for $V_h(1)$, the expected value of rejecting the $h^{\mathrm{th}}$ candidate is $\mathbb{E}[V_{h+1}(s_{h+1})]$.


We start by calculating the optimal value function at the last stage $n$, i.e., the values of $V_n(s_n)$ for all $s_n\in\{1,...,n\}$. $V_n(s_n)=1$ if $s_n=1$ and $V_n(s_n)=0$ otherwise because at stage $n$ we have interviewed all of the $n$ candidates and $s_n=1$ means that the last candidate is the best among $n$ candidates. Then we can calculate $V_h(s_h)$ for $h=n-1,\ldots,1$ using backward search with the Bellman equation above. After obtaining the value function, we can determine the optimal policy by a forward pass.


## Codes

### Backward Search

We can do backward search based on the the Bellman equation and the values of $V_n(s_n)$.

For the Python function `backward_cal` in the next cell, the input of the function is $n$. The output `value_function` is the value function $V_{h}(s_{h})$ for all $h\in\{1,\ldots,n\}$ and $s_h\in\{1,\ldots,h\}$. Note that we initialize the `value_function` such that its shape is $(n+1, n+1)$ and `value_function[h,s_h]` means the value function at stage $h$ and state $s_h$.


```python
# Import packages. Run this cell.

import numpy as np
import matplotlib.pyplot as plt
```

<!-- grade_id: question_3_code, points: 0, grade: false, solution: true -->
```python
def backward_cal(n):
    """
    Calculate the optimal value function $V_h(s_h)$ using backward seach
    Args:
        n: the total number of stages, n >= 2
    Returns:
        value_function: a numpy array with shape (n+1, n+1). value_function[h, s_h] represents $V_{h}(s_h)$.
    """
    value_function = np.zeros((n + 1, n + 1))
    value_function[n, 1] = 1
    stage = n - 1
    while stage >= 1:
        for state in range(2, n + 1):
            value_function[stage, state] = (value_function[stage + 1, 1] + value_function[stage + 1, 2] * stage) / (stage + 1)
        value_function[stage, 1] = max(stage / n, value_function[stage, 2])
        stage = stage - 1
  
    return value_function

```

<!-- grade_id: question_3_sample_test, points: 0, grade: true, solution: false -->
```python
# Sample Test, checking the output of the function backward_cal

# Sample input
n = 3

# Sample output
value_function = np.array([[ 0.0, 0.0, 0.0, 0.0],
                           [ 0.0, 1/2, 0.0, 0.0],
                           [ 0.0, 2/3, 1/3, 0.0],
                           [ 0.0, 1.0, 0.0, 0.0]])

# Sample test
func_out = backward_cal(n)
for h in range(1, n + 1):
    for s_h in range(1, h + 1):
        assert round(func_out[h, s_h], 4) == round(value_function[h, s_h], 4), "The sample test failed."

```

### Closed-Form Solution

Consider $n\ge 3$. Let $h^*$ be such that 
$\frac{1}{h^*}+\frac{1}{h^*+1}+...+\frac{1}{n-1} \le 1 < \frac{1}{h^*-1}+\frac{1}{h^*}+...+\frac{1}{n-1}$.

For example, if $n=10$, then $h^*=4$.

Given $n$, $h$, and $s_h$, the value function $V_h(s_h)$ can also be calculated using the closed-form equations [1] below.

For all $h^*-1 \le h \le n-1$,

\begin{equation}
    \begin{aligned}
        V_h(1) =& \max\left\{\frac{h}{n}, \frac{h}{n}\left(\frac{1}{h}+\frac{1}{h+1}+...+\frac{1}{n-1}\right)\right\}\\
        V_h(s_h) =& \frac{h}{n}\left(\frac{1}{h}+\frac{1}{h+1}+...+\frac{1}{n-1}\right), s_h=2,...,h
    \end{aligned}
\end{equation}

and for all $1 \le h\le h^*-1$,

\begin{align}
    V_h(s_h) =& \frac{h^*-1}{n}\left(\frac{1}{h^*-1}+\frac{1}{h^*}+...+\frac{1}{n-1}\right), \forall s_h=1,...,h.
\end{align}
    
We implemented this closed-form calculation as the Python function `V_h_closed_form`. We can verify these closed-form equations by choosing several sets of $n$, $h$, and $s_h$ as inputs and then compare the values of $V_h(s_h)$ obtained by the Python function `backward_cal` with those obtained by `V_h_closed_form`.

```python
def V_h_closed_form(n, h, s_h):
    """
    Calculate the optimal value function $V_h(s_h)$ using the closed-form equations
    Args:
        n: the total number of stages
        h: stage h
        s_h: the state s_h at stage h
    Returns:
        v_h_s_h: the value of the optimal value function $V_h(s_h)$
    """
    if h == n:
        if s_h == 1:
            v_h_s_h = 1
        else:
            v_h_s_h = 0
    else:
        for h_star in range(2, n):
            left = sum([1 / i for i in range(h_star, n)])
            right = left + 1 / (h_star - 1)
            if left <= 1 and right > 1:
                break
        if h >= h_star:
            if s_h == 1:
                v_h_s_h = max(h / n, h / n * sum([1 / i for i in range(h, n)]))
            else:
                v_h_s_h = h / n * sum([1 / i for i in range(h, n)])
        else:
            v_h_s_h = (h_star - 1) / n * sum([1 / i for i in range(h_star - 1, n)])
    
    return v_h_s_h
```

<!-- grade_id: question_4_code, points: 0, grade: true, solution: true -->
```python
# Comparison
n = 20  # any positive integer you like
h = 9  # any integer between 1 and n
s_h = 2  # any integer between 1 and h
V_h = backward_cal(n)
print(V_h[h, s_h], V_h_closed_form(n, h, s_h))
h = 11  # any integer between 1 and n
s_h = 2  # any integer between 1 and h
print(V_h[h, s_h], V_h_closed_form(n, h, s_h))
h = 11  # any integer between 1 and n
s_h = 1  # any integer between 1 and h
print(V_h[h, s_h], V_h_closed_form(n, h, s_h))
```

Output:
```
0.3734471314289426 0.37344713142894265
0.3403242717464854 0.3403242717464854
0.55 0.55
```

### Find the Optimal Policy

The following function `optimal_action` will output the optimal action `a_h` given the inputs of the value function `value_function`, the number of stages `n`, stage `h` and state `s_h`.

<!-- grade_id: question_5_code, points: 0, grade: true, solution: true -->
```python
def optimal_action(value_function, n, h, s_h):
    """
    Calculate the optimal action $a_h$ at stage $h$
    Args:
        value_function: a numpy array with shape (n+1, n+1). value_function[h, s_h] represents $V_{h}(s_h)$.
        n: the total number of stages, n >= 2
        h: stage h
        s_h: the state s_h at stage h
    Returns:
        a_h: the optimal action, 1 means hiring the candidate, 0 means rejecting the candidate
    """
    a_h = None
    if h == n:
        a_h = 1
    else:
        if s_h > 1:
            a_h = 0
        else:
            if h/n > (value_function[h + 1, 1] + value_function[h + 1, 2] * h) / (h + 1):
                a_h = 1
            else:
                a_h = 0
    return a_h

```

```python
# Sample Test, checking the output of the function optimal_action

# Sample input
n = 3
value_function = backward_cal(n)
h = 1
s_h = 1

# Sample output
a_h = 0

# Sample test
func_out = optimal_action(value_function, n, h, s_h)
assert func_out == a_h, "The sample test failed."
```

### Optimal Policy in Closed-Form

Using the closed-form solution of the value function and the Bellman equation, we can obtain the optimal policy, that is, at stage $h$, hire the candidate if

- $h\ge h^*$, and
- they are the best candidate so far ($s_h=1$).

In other words, we reject first $h^*-1$ candidates and hire the best candidate so far after that. For large $n$, $h^* \approx \frac{n}{e}$ [1].


## Reference
[1] MJ Beckmann. Dynamic programming and the secretary problem. Computers & Mathematics with Applications, 19(11):25–28, 1990.