# Chapter 1 Dynamic Programming and Bellman Equations


## Example: Secretary Problem

In this example, we will solve the secretary problem using stochastic dynamic programming (DP).

### Secretary Problem

The secretary problem [1] can be described as follows. Imagine an administrator who wants to hire the best secretary out of $n$ available candidates. They are interviewed in succession in random order, that is, for each interview the administrator picks a candidate uniformly at random from the remaining candidates. After interviewing the $k^{\mathrm{th}}$ candidate, the administrator knows the ranking of the $k^{\mathrm{th}}$ candidate among the candidates interviewed so far, but is unaware of the quality of unseen candidates. The administrator needs to determine whether to hire or reject the $k^{\mathrm{th}}$ candidate immediately after the interview. Note that the administrator cannot go back to hire candidate $1,\ldots,k-1$ after interviewing the $k^{\mathrm{th}}$ candidate. At most one candidate can be hired. The goal is to maximize the probability of hiring the best candidate.

### Problem Formulation

We can formulate the secretary problem into a finite-horizon stochastic DP problem as follows:

- Stage $k$: the time when the administrator makes decision after interviewing the $k^{\mathrm{th}}$ candidate. $k\in\{1,...,n\}$.
    
- State $x_k$: the ranking of the $k^{\mathrm{th}}$ candidate among the candidates interviewed so far. $x_k\in\{1,...,k\}$.

    Define a terminal state $g$ such that if the administrator has already hired one candidate in one of the previous interviews, then $x_{k}=g, x_{k+1}=g,\ldots,x_{n}=g$. That is, once the administrator decides to hire a candidate, the process will go to the terminal state and remain in the terminal state.

- Action $u_k$:
    
  If $x_k\neq g$, then

  - $u_k=1$ means hiring the $k^{\mathrm{th}}$ candidate;
  - $u_k=0$ means rejecting the $k^{\mathrm{th}}$ candidate.

  If $x_k=g$, then the only available action is $u_k=0$, meaning rejecting the $k^{\mathrm{th}}$ candidate, because $x_k=g$ means that the administrator has already hired one candidate previously.
    
- Reward $r_k(x_k,u_k)$: 
    
    $r_k(x_k,u_k)=1$ if $x_k\neq g$, $u_k=1$ and the $k^{\mathrm{th}}$ candidate is the best candidate among all the $n$ candidates. 
    
    Otherwise, $r_k(x_k,u_k)=0$.
    
    Note that if $x_k\neq 1$, $r_k(x_k,1)=0$ because $x_k\neq 1$ means that the $k^{\mathrm{th}}$ candidate is not the best among the $k$ candidates that have been interviewed and of course not the best among all candidates. If $x_k=1$, we can calculate the probability that the $k^{\mathrm{th}}$ candidate is the best among all candidates as follows:

    \begin{align*}
        \Pr\left(\mbox{the $k^{\mathrm{th}}$ candidate is the best among all candidates}|x_k=1\right)=\frac{{n-1 \choose k-1} (k-1)!(n-k)!}{{n \choose k} (k-1)!(n-k)!}=\frac{k}{n}.
    \end{align*}

    Therefore, $r_k(1,1)=1$ with probability $k/n$.
    
- Transition probabilities:
    
    For any $m\in\{1,\ldots,k+1\}$ and $j\in\{1,\ldots,k\}$, we have

    \begin{align*}
        \Pr(x_{k+1}=m|x_k=j,u_k=0)=\frac{{n \choose k+1} (k-1)!(n-k-1)!}{{n \choose k+1} {k+1 \choose k}(k-1)!(n-k-1)!}=\frac{1}{k+1},
    \end{align*}

    which means that $x_{k+1}$ will be uniformly distributed in $\{1,\ldots,k+1\}$ if $x_k\neq g$ and $u_k=0$. And for any $j\in\{1,\ldots,k\}$,

    \begin{align*}
        &\Pr(x_{k+1}=g|x_k=j,u_k=0)=0,\\
        &\Pr(x_{k+1}=g|x_k=j,u_k=1)=1,\\
        &\Pr(x_{k+1}=g|x_k=g,u_k=0)=1.
    \end{align*}
    
- Goal: maximize the probability of hiring the best candidate. Note that this goal can be written as 
    \begin{align*}
        \max_{\mu_1,\ldots,\mu_n} \mathbb{E} \left[\sum_{k=1}^{n} r_k(x_k,\mu_k(x_k))\right].
    \end{align*}

### Value Function

Let $V_k(x_k)$ denote the optimal value function for state $x_k$ at stage $k$, defined by
$$
V_k(x_k) = \max_{\mu_k,\ldots,\mu_n} \mathbb{E} \left[\sum_{j=k}^n r_j(x_j, \mu_j(x_j))\right].
$$
The optimal value function $V_k(x_k)$ can be interpreted as the probability that the best candidate will be hired under the optimal policy given $x_k$ at stage $k$. 

For example, let $n = 10$, $k = 5$, $x_5 = 1$. In this scenario, there are 10 candidates total. After interviewing the $5^{\mathrm{th}}$ candidate, we know that this candidate ranks first among the 5 candidates that have been interviewed. Then $V_5(1)$ is the probability that we can successfully hire the best candidate if we follow the optimal policies at and after stage $5$.

### Bellman Equation

Based on the above definition and analysis, the Bellman equation for $k=1,...,n-1$ can be written as:

\begin{equation}
    \begin{aligned}
        V_k(1) =& \max \left\{\mathbb{E}[r_k(1,1)], \mathbb{E}[V_{k+1}(x_{k+1})]\right\} = \max\left\{\frac{k}{n},\frac{1}{k+1}\sum_{x=1}^{k+1}V_{k+1}(x)\right\}\\
        V_k(x_{k}) =& \max \left\{0, \mathbb{E}[V_{k+1}(x_{k+1})]\right\} = \frac{1}{k+1}\sum_{x=1}^{k+1}V_{k+1}(x), x_k=2,...,k.
    \end{aligned}
\end{equation}

For $V_k(1)$, the first term inside the maximum function, $\mathbb{E}[r_k(1,1)]$, is the expected reward given $u_k=1$, i.e., hiring the $k^{\mathrm{th}}$ candidate. The total future reward is 0 since the process terminates. The second term $\mathbb{E}[V_{k+1}(x_{k+1})]$ is the expected value given $u_k=0$, i.e., rejecting the $k^{\mathrm{th}}$ candidate, since the instantaneous reward is 0.

For $V_k(x_{k})$ where $x_k=2,...,k$, the expected value for hiring the $k^{\mathrm{th}}$ candidate is 0 since the ranking of the $k^{\mathrm{th}}$ candidate among all the $n$ candidates must be larger than or equal to 2. Similar to the case for $V_k(1)$, the expected value of rejecting the $k^{\mathrm{th}}$ candidate is $\mathbb{E}[V_{k+1}(x_{k+1})]$.


We start by calculating the optimal value function at the last stage $n$, i.e., the values of $V_n(x_n)$ for all $x_n\in\{1,...,n\}$. $V_n(x_n)=1$ if $x_n=1$ and $V_n(x_n)=0$ otherwise because at stage $n$ we have interviewed all of the $n$ candidates and $x_n=1$ means that the last candidate is the best among $n$ candidates. Then we can calculate $V_k(x_k)$ for $k=n-1,\ldots,1$ using backward search with the Bellman equation above. After obtaining the value function, we can determine the optimal policy by a forward pass.


## Codes

### Backward Search

We can do backward search based on the the Bellman equation and the values of $V_n(x_n)$.

For the Python function `backward_cal` in the next cell, the input of the function is $n$. The output `value_function` is the value function $V_{k}(x_{k})$ for all $k\in\{1,\ldots,n\}$ and $x_k\in\{1,\ldots,k\}$. Note that we initialize the `value_function` such that its shape is $(n+1, n+1)$ and `value_function[k,x_k]` means the value function at stage $k$ and state $x_k$.


```python
# Import packages. Run this cell.

import numpy as np
import matplotlib.pyplot as plt
```

<!-- grade_id: question_3_code, points: 0, grade: false, solution: true -->
```python
def backward_cal(n):
    """
    Calculate the optimal value function $V_k(x_k)$ using backward seach
    Args:
        n: the total number of stages, n >= 2
    Returns:
        value_function: a numpy array with shape (n+1, n+1). value_function[k, x_k] represents $V_{k}(x_k)$.
    """
    value_function = np.zeros((n + 1, n + 1))
    value_function[n, 1] = 1
    stage = n - 1
    while stage >= 1:
        for state in range(2, n + 1):
            value_function[stage, state] = (value_function[stage + 1, 1] + value_function[stage + 1, 2] * stage) / (stage + 1)
        value_function[stage, 1] = max(stage / n, value_function[stage, 2])
        stage = stage - 1
  
    return value_function

```

<!-- grade_id: question_3_sample_test, points: 0, grade: true, solution: false -->
```python
# Sample Test, checking the output of the function backward_cal

# Sample input
n = 3

# Sample output
value_function = np.array([[ 0.0, 0.0, 0.0, 0.0],
                           [ 0.0, 1/2, 0.0, 0.0],
                           [ 0.0, 2/3, 1/3, 0.0],
                           [ 0.0, 1.0, 0.0, 0.0]])

# Sample test
func_out = backward_cal(n)
for k in range(1, n + 1):
    for x_k in range(1, k + 1):
        assert round(func_out[k, x_k], 4) == round(value_function[k, x_k], 4), "The sample test failed."

```

### Closed-Form Solution

Consider $n\ge 3$. Let $k^*$ be such that 
$\frac{1}{k^*}+\frac{1}{k^*+1}+...+\frac{1}{n-1} \le 1 < \frac{1}{k^*-1}+\frac{1}{k^*}+...+\frac{1}{n-1}$.

For example, if $n=10$, then $k^*=4$.

Given $n$, $k$, and $x_k$, the value function $V_k(x_k)$ can also be calculated using the closed-form equations [1] below.

For all $k^*-1 \le k \le n-1$,

\begin{equation}
    \begin{aligned}
        V_k(1) =& \max\left\{\frac{k}{n}, \frac{k}{n}\left(\frac{1}{k}+\frac{1}{k+1}+...+\frac{1}{n-1}\right)\right\}\\
        V_k(x_k) =& \frac{k}{n}\left(\frac{1}{k}+\frac{1}{k+1}+...+\frac{1}{n-1}\right), x_k=2,...,k,
    \end{aligned}
\end{equation}

and for all $k$ such that $1 \le k\le k^*-1$,

\begin{align*}
    V_k(x_k) =& \frac{k^*-1}{n}\left(\frac{1}{k^*-1}+\frac{1}{k^*}+...+\frac{1}{n-1}\right), \forall x_k=1,...,k.
\end{align*}
    
We implemented this closed-form calculation as the Python function `V_k_closed_form`. We can verify these closed-form equations by choosing several sets of $n$, $k$, and $x_k$ as inputs and then compare the values of $V_k(x_k)$ obtained by the Python function `backward_cal` with those obtained by `V_k_closed_form`.

```python
def V_k_closed_form(n, k, x_k):
    """
    Calculate the optimal value function $V_k(x_k)$ using the closed-form equations
    Args:
        n: the total number of stages
        k: stage k
        x_k: the state x_k at stage k
    Returns:
        v_k_x_k: the value of the optimal value function $V_k(x_k)$
    """
    if k == n:
        if x_k == 1:
            v_k_x_k = 1
        else:
            v_k_x_k = 0
    else:
        for k_star in range(2, n):
            left = sum([1 / i for i in range(k_star, n)])
            right = left + 1 / (k_star - 1)
            if left <= 1 and right > 1:
                break
        if k >= k_star:
            if x_k == 1:
                v_k_x_k = max(k / n, k / n * sum([1 / i for i in range(k, n)]))
            else:
                v_k_x_k = k / n * sum([1 / i for i in range(k, n)])
        else:
            v_k_x_k = (k_star - 1) / n * sum([1 / i for i in range(k_star - 1, n)])
    
    return v_k_x_k
```

<!-- grade_id: question_4_code, points: 0, grade: true, solution: true -->
```python
# Comparison
n = 20  # any positive integer you like
k = 9  # any integer between 1 and n
x_k = 2  # any integer between 1 and h
V_k = backward_cal(n)
print(V_k[k, x_k], V_k_closed_form(n, k, x_k))
k = 11  # any integer between 1 and n
x_k = 2  # any integer between 1 and h
print(V_k[k, x_k], V_k_closed_form(n, k, x_k))
k = 11  # any integer between 1 and n
x_k = 1  # any integer between 1 and h
print(V_k[k, x_k], V_k_closed_form(n, k, x_k))
```

### Find the Optimal Policy

The following function `optimal_action` will output the optimal action `u_k` given the inputs of the value function `value_function`, the number of stages `n`, stage `k` and state `x_k`.

<!-- grade_id: question_5_code, points: 0, grade: true, solution: true -->
```python
def optimal_action(value_function, n, k, x_k):
    """
    Calculate the optimal action $u_k$ at stage $k$
    Args:
        value_function: a numpy array with shape (n+1, n+1). value_function[k, x_k] represents $V_{k}(x_k)$.
        n: the total number of stages, n >= 2
        k: stage k
        x_k: the state x_k at stage k
    Returns:
        u_k: the optimal action, 1 means hiring the candidate, 0 means rejecting the candidate
    """
    u_k = None
    if k == n:
        u_k = 1
    else:
        if x_k > 1:
            u_k = 0
        else:
            if k/n > (value_function[k + 1, 1] + value_function[k + 1, 2] * k) / (k + 1):
                u_k = 1
            else:
                u_k = 0
    return u_k

```

```python
# Sample Test, checking the output of the function optimal_action

# Sample input
n = 3
value_function = np.array([[ 0.0, 0.0, 0.0, 0.0],
                           [ 0.0, 1/2, 0.0, 0.0],
                           [ 0.0, 2/3, 1/3, 0.0],
                           [ 0.0, 1.0, 0.0, 0.0]])
k = 1
x_k = 1

# Sample output
u_k = 0

# Sample test
func_out = optimal_action(value_function, n, k, x_k)
assert func_out == u_k, "The sample test failed."
```

### Optimal Policy in Closed-Form

Using the closed-form solution of the value function and the Bellman equation, we can obtain the optimal policy, that is, at stage $k$, hire the candidate if

- $k\ge k^*$, and
- they are the best candidate so far ($x_k=1$).

In other words, we reject first $k^*-1$ candidates and hire the best candidate so far after that. For large $n$, $k^* \approx \frac{n}{e}$ [1].


## Reference
[1] MJ Beckmann. Dynamic programming and the secretary problem. Computers & Mathematics with Applications, 19(11):25–28, 1990.